# HUE 05 (Ausgabe 26.11.2025 15:45 - Abgabe Mi. 3.12.2025 11:00)

Für fast alles gibt es in Python bereits eingebaute Funktionen (z.B. `count`,`sum`,`min`,`index`,...).    
Warum sollen Sie diese eingebauten Funktionen NICHT benutzen?    
Hier geht es darum, an einfachen Beispielaufgaben wesentliche Programmiertechniken zu lernen. 
Deshalb sollen diese Funktionen NICHT benutzt werden. 

## Aufgabe 1: Rotation nach links

Schreiben Sie eine Funktion `rot`, die ein Tupel (oder einen String) als Parameter bekommt, und ein Tupel (oder einen String) zurückgibt, in der die Elemente (oder die Zeichen) um eine Position nach links rotiert sind, d.h., das erste Element (oder das erste Zeichen) kommt ans Ende des Tupels (oder des Strings).   
Beispiele: 

Beispielsweise soll der Aufruf
```python
rot("abcd")
``` 
sollte den Wert `'bcda'` zurückgeben und
```python
rot(((1,2),(3,4),5))
```
den Wert  `((3, 4), 5, (1, 2))`

<!-- grade_id: question1_ans, points: 0, grade: false, solution: true -->
```python
# Implementieren Sie hier Ihre Funktion

def rot(s):
    if len(s) <= 1:
        return s
    return s[1:]+s[:1]
```

<!-- grade_id: question1_test, points: 3, grade: true, solution: false -->
```python
# Tests für `rot`
assert rot("abcd") == "bcda", 'public test failed: rot("abcd") == "bcda"'
assert rot(((1, 2), (3, 4), 5)) == ((3, 4), 5, (1, 2)), "rot should also work for tuples"
```

```python
# Führen Sie nach der Zelle mit Ihrem Programm diese Zelle mit verschiedenen Aktualparametern aus, 
# und überprüfen Sie, ob der Aufruf Ihrer Funktion das gewünschte Ergebnis liefert. 


rot("abcd")
```

Output:
```
'bcda'
```

## Aufgabe 2: Finde drittkleinstes Element einer Liste (bei gegebenen verschieden großen Elementen)

Schreiben Sie eine Funktion `third_smallest`, die das drittkleinste Element einer Liste zurückgibt.   
Setzen Sie voraus, dass alle Elemente verschieden groß sind und die Liste mindestens 3 Elemente hat.  
Programmieren Sie ein Lösung, die das drittkleinste Element bestimmt, ohne die Listenelemente zu sortieren.

Sie können davon ausgehen, dass alle Elemente in der Liste kleiner als `float('inf')` sind.

Beispielsweise soll der Aufruf
```python
third_smallest([3,1,4,2,1.5,5])
```
den Wert `2` zurückgeben.

<!-- grade_id: question2_ans, points: 0, grade: false, solution: true -->
```python
# Implementieren Sie hier Ihre Funktion

def third_smallest(values):
    """Gibt das drittkleinste Element der Liste zurück."""
    # YOUR CODE HERE
    values.sort()
    return values[2]
```

<!-- grade_id: question2_test, points: 6, grade: true, solution: false -->
```python
# Tests for `third_smallest`
assert third_smallest([3, 1, 4, 2, 1.5, 5]) == 2, "third_smallest should return 2"
assert third_smallest([10, 0, -5, 7]) == 7, "third_smallest should return 7"
```

```python
# Führen Sie nach der Zelle mit Ihrem Programm diese Zelle mit verschiedenen Aktualparametern aus, 
# und überprüfen Sie, ob der Aufruf Ihrer Funktion das gewünschte Ergebnis liefert. 

third_smallest([4.1,4,3,7,1.5,5])
```

Output:
```
4
```

## Aufgabe 3: Flächeninhalt von Polygonen 

Schreiben Sie eine Funktion `flächeninhalt`, welche als Parameter ein Polygon in Form eines Tupels an Punkten $(x,y)$ übergeben bekommt. Die Funktion soll den Flächeninhalt des übergebenen Polygons berechnen und zurückgeben. Ein Rechteck mit Breite 1 und Höhe 2, dessen untere linke Ecke im Punkt $(0,0)$ liegt, könnte zum Beispiel so wie in Zelle 1 gegeben sein. Sie können davon ausgehen, dass die Kanten kreuzungsfrei sind. 

Wenn
$$
(x_0,y_0) , (x_1,y_1), \dots, (x_n,y_n), (x_0,y_0)
$$
die Eckpunkte des Polygons sind (der $n+1$-te Punkt ist gleich 
dem $0$-ten), dann lautet die Formel für die Fläche: 
$$
    \frac{x_0 y_1 - y_0 x_1}{2} + \frac{x_1 y_2 - y_1 x_2}{2} + \frac{x_2 y_3 - y_2 x_3}{2} + \dots + \frac{x_n y_0 - x_0 y_n}{2}
    %% (X0*Y1 - Y0*X1)/2 + (X1*Y2 - Y1*X2)/2 + (X2*Y3 - Y2*X3)/2 + ... + (Xn*Y0 - X0*Yn)/2 
$$

Beispielsweise soll der Aufruf

```python
flächeninhalt(((0, 0), (1, 0), (1, 2), (0, 2)))
```
den Wert `2.0` zurückgeben.

<!-- grade_id: question3_ans, points: 0, grade: false, solution: true -->
```python
# Implementieren Sie hier Ihre Funktion

def flächeninhalt(polygon):
    """Berechnet den Flächeninhalt eines Polygons mittels Shoelace-Formel."""
    # YOUR CODE HERE
    fläche=0
    for i in range(len(polygon)):
        output=0
        if i == range(len(polygon)):
            x0,y0=polygon[i]
            x1,y1=polygon[0]
            fläche+=(x0*y1-y0*x1)/2
        else:
            x0,y0=polygon[i-1]
            x1,y1=polygon[i]
            fläche+=(x0*y1-y0*x1)/2
    return fläche
    raise NotImplementedError()
```

<!-- grade_id: question3_test, points: 10, grade: true, solution: false -->
```python
# Tests for `flächeninhalt`
assert abs(flächeninhalt(((0, 0), (2, 0), (2, 2), (0, 2))) - 4.0) < 1e-9, "Square 2x2 should have area 4"
assert abs(flächeninhalt(((0, 0), (1, 0), (1, 2), (0, 2))) - 2.0) < 1e-9, "Rectangle 1x2 should have area 2"
```

```python
# Testen Sie Ihre Funktion flächeninhalt mit verschiedenen Aktualparametern
polygon = ((0, 0), (2, 0), (2, 1), (0, 1))
flächeninhalt(polygon)
```

Output:
```
2.0
```

## Aufgabe 4: Duplikate in einer Liste von rechts entfernen (ohne vordefinierte Funktion zu nutzen) 
Programmieren Sie eine Funktion `remove_duplicates`, die eine Liste `L` als Parameter übergeben bekommt und, von rechts ausgehend, nach Duplikaten in der Liste sucht und alle Duplikate in der Liste löscht.    
Gemeint ist, dass bei Duplikaten das am weitesten rechts in der Liste stehende Element erhalten bleiben soll 
und die weiter links stehenden gleichen Elemente gelöscht werden sollen.    

Des Weiteren sollen Sie KEINE vordefinierten Funktionen verwenden, um diese Aufgabe zu lösen.

**Wichtig**: Es soll keine neue Liste ausgegeben werden, sondern die als Parameter übergebene Liste soll geändert werden. 

Beispielsweise soll nach dem Aufruf
```python
L = [4, 1, 4, 2, 10, 8, 2, 4, 5, 6, 8]
remove_duplicates(L)
```
gelten, dass
```python
L == [1, 10, 2, 4, 5, 6, 8]
```

<!-- grade_id: question4_ans, points: 0, grade: false, solution: true -->
```python
# Implementieren Sie hier Ihre Funktion

def remove_duplicates(L):
    # YOUR CODE HERE
    i=len(L)-1
    while i>0:
        j=i-1
        while j>=0:
            if L[j] == L[i]:
                del L[j]
                i-=1
            j-=1
        i-=1
    return L
```

<!-- grade_id: question4_test, points: 0, grade: true, solution: false -->
```python
# Tests for `remove_duplicates`
data = [4, 1, 4, 2, 10, 8, 2, 4, 5, 6, 8]
remove_duplicates(data)
assert data == [1, 10, 2, 4, 5, 6, 8], "List should contain only rightmost duplicates"
```

```python
# Testen Sie Ihre Funktion mit verschiedenen Aktualparamtern
L = [4, 1, 4, 2, 10, 8, 2, 4, 5, 6, 8]
remove_duplicates(L)
L
```

Output:
```
[1, 10, 2, 4, 5, 6, 8]
```

## Aufgabe 5: Ringliste simulieren Abzählvers 
Schreiben Sie eine Funktion `abzählen`, welche eine Liste `L` und einen Integer `n` übergeben bekommt. Die Funktion soll einen Abzählvers mit `n` Silben simulieren. Beispielsweise würde bei dem Vers "Du fliegst raus" jedes dritte Element aus der Liste entfernt werden. Ihre Funktion soll also jedes `n`-te Element aus der Liste entfernen und ausgeben. Nach dem Entfernen eines Elements wird weiter abgezählt, bis die Liste leer ist. Wenn man am Ende der Liste angekommen ist, wird von vorne weitergezählt. Die Funktion soll eine andere Liste ausgeben, in welcher die Elemente von `L` in der Reihenfolge gespeichert sind, in welcher sie aus `L` entfernt wurden.

Beispielsweise soll der Aufruf
```python
abzählen(["anna","bernd","christine","david","emily","frank","gesa"], 3)
```
die Liste 
```python
['christine', 'frank', 'bernd', 'gesa', 'emily', 'anna', 'david']
```
zurückgeben. 

Ihr Programm soll keinen Aufruf von `pop` verwenden.

<!-- grade_id: question5_ans, points: 0, grade: false, solution: true -->
```python
# Implementieren Sie hier Ihre Funktion

def abzählen(L, n):
    # YOUR CODE HERE
    new_list=[]
    a=len(L)
    i=0
    while len(L) > 0:
        i=(i+n-1)%len(L)
        print(L, len(L), i)
        new_list.append(L[i])
        del L[i]
    return new_list
abzählen(["anna","bernd","christine","david","emily","frank","gesa"], 3)
```

Output:
```
['anna', 'bernd', 'christine', 'david', 'emily', 'frank', 'gesa'] 7 2
['anna', 'bernd', 'david', 'emily', 'frank', 'gesa'] 6 4
['anna', 'bernd', 'david', 'emily', 'gesa'] 5 1
['anna', 'david', 'emily', 'gesa'] 4 3
['anna', 'david', 'emily'] 3 2
['anna', 'david'] 2 0
['david'] 1 0
['christine', 'frank', 'bernd', 'gesa', 'emily', 'anna', 'david']
```

<!-- grade_id: question5_test, points: 10, grade: true, solution: false -->
```python
# Tests for `abzählen`
assert abzählen(["anna", "bernd", "christine", "david", "emily", "frank", "gesa"], 3) == ['christine', 'frank', 'bernd', 'gesa', 'emily', 'anna', 'david'], "Counting rhyme failed for step size 3"
assert abzählen(["a", "b", "c"], 1) == ['a', 'b', 'c'], "Step size 1 should remove items in order"
```

```python
# Testen Sie Ihre Funktion mit verschiedenen Aktualparamtern
L = ["anna","bernd","christine","david","emily","frank","gesa"]
abzählen(L, 3)
```

```python

```